Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tejas Yeshwantrao Wandile"
COLLABORATORS = ""
```

---

<!-- grade_id: cell-29d848ad7ec80b15, points: 0, grade: false, solution: true -->
```python
# Exercise 1
def is_palindrome(input_string):
    # Complete this function to determine if the input_string is a palindrome. Return True or False
    # Use square brackets to reverse the input_string! Make sure to lower the input string before testing!
    
    lower_str = input_string.lower()
    reverse_str = lower_str[::-1]
    if lower_str == reverse_str:
        return True
    else:
        return False
```

<!-- grade_id: cell-20dd267f296c8074, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 1
assert is_palindrome("Kayak") == True
assert is_palindrome("Rotator") == True
assert is_palindrome("AACA") == False

```

<!-- grade_id: cell-011ea97b9fa805a5, points: 3, grade: true, solution: false -->
```python
# Test cell  for exercise 1 (hidden)
```

<!-- grade_id: cell-d344708f226ea451, points: 0, grade: false, solution: true -->
```python
# Exercise 2: Season from Month and Day
def determine_season(month, day):

    # The year is divided into four season: spring, summer, fall (or autumn) and winter.
    # While the exact dates that the seasons change vary a little bit from year to
    # year because of the way that the calender is constructed, we will use the following
    # dates for this exercise:

    # Season  -- First Day
    # Spring  -- March 20
    # Summer  -- June 21
    # Fall  -- September 22
    # Winter    -- December 21

    # Complete this function which takes in as inputs a month and day. It should
    # output the season.
    # input 1: month -- str
    # input 2: day -- int

    # output: month -- str (Spring, Summer, Fall, Winter)
    import datetime
    months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December']
    index = months.index(month)+1
    d1 = datetime.datetime(2018,index, day)
    March20 = datetime.datetime(2018, 3, 20)
    June21 = datetime.datetime(2018, 6, 21)
    September22 = datetime.datetime(2018, 9, 22)
    December21 = datetime.datetime(2018, 12, 21)
    if d1 < March20:
        return "Winter"
    elif d1 >= March20 and d1< June21:
        return "Spring"
    elif d1 >= June21 and d1 < September22:
        return "Summer"
    elif d1 >= September22 and d1 < December21:
        return "Fall"
    #print(index)
    print(d1)
```

<!-- grade_id: cell-264aab0e7ab82be4, points: 4, grade: true, solution: false -->
```python
# Test cell  for exercise 2
assert determine_season('March', 21) == 'Spring'
assert determine_season('June', 21) == 'Summer'
assert determine_season('November', 21) == 'Fall'
assert determine_season('January', 21) == 'Winter'
```

<!-- grade_id: cell-5aa9eaa720eeab1e, points: 4, grade: true, solution: false -->
```python
# Test cell  for exercise 2 (hidden)
```

<!-- grade_id: cell-f8bbb086e9047977, points: 0, grade: false, solution: true -->
```python
#Exercise 3: Check a password

def check_password(password):

    # In this exercise you will complete this function to determine whether or not
    # a password is good. We will define a good password to be a one that is at least
    # 8 characters long and contains at least one uppercase letter, at least one lowercase
    # letter, and at least one number. This function should return True if the password
    # passed to it as its only parameter is good. Otherwise it should return False. 
    #
    # input: password (str)
    # output: True or False (bool)
    
    def checkPassword(password):
        lower = False
        upper = False
        length = False
        digit = False
        for ch in password:
            if ch.isupper():
                upper = True
            if ch.islower():
                lower = True
            if ch.isdigit():
                digit = True
                
        if len(password) >= 8:
            length = True
            
        if(upper and lower and length and digit):
            return True
         
        return False

    

    if checkPassword(password):
        return True
    return False



```

<!-- grade_id: cell-1b60abed8b8bc04b, points: 4, grade: true, solution: false -->
```python
# Test cell  for exercise 3
assert check_password('test1234') == False
assert check_password('password123') == False
assert check_password('SuperPasswrd90') == True
assert check_password('letmein!') == False

```

<!-- grade_id: cell-7eaad572678f3ec3, points: 4, grade: true, solution: false -->
```python
# Test cell  for exercise 3

```

<!-- grade_id: cell-a2d16a8840027f58, points: 0, grade: false, solution: true -->
```python
# Exercise 4
def average_word_len_in_sentence(sentence):
    # Complete this function to calculate the average
    # word length in a sentence
    # Input: sentence
    # Output: average word length in sentence
    # Hint: count punctuations with whatever word they are `touching`
    # Hint: round the average to two decimal places
    
    words = sentence.split();
    count = 0
    for ch in words:
        count += len(ch)
    return round(count/len(words),2)
    
```

```python
average_word_len_in_sentence("This is a test sentence!")
```

Output:
```
4.0
```

<!-- grade_id: cell-b0fb7835c133ea39, points: 1, grade: true, solution: false -->
```python
# Text cell for Exercise 4 
assert average_word_len_in_sentence('This is a test sentence!') == 4
```

Output:
```
4.0
AssertionError                            Traceback (most recent call last)
/tmp/ipykernel_346799/2466448239.py in <module>
      1 # Text cell for Exercise 4
----> 2 assert average_word_len_in_sentence('This is a test sentence!') == 4

AssertionError:
```

<!-- grade_id: cell-f4ef1dd8731cf063, points: 1, grade: true, solution: false -->
```python
# Text cell for Exercise 4(hidden)
```

<!-- grade_id: cell-8530a32cd73998fd, points: 0, grade: false, solution: true -->
```python
#Exercise 5
def wc(filename):
    # Complete this function to count the number of lines, words, and chars in a file. 
    # Input: filename
    # Output: a tuple with line count, word count, and char count -- in this order


    count_lines = 0
    count_words = 0
    count_chars = 0

    with open(filename, 'r') as f:
        for line in f:
            words = line.split()

            count_lines += 1
            count_words += len(words)
            count_chars += len(line)
            
    return (count_lines, count_words, count_chars)

```

<!-- grade_id: cell-ec3e9bad05323884, points: 4, grade: true, solution: false -->
```python
# Text cell for Exercise 5 
assert wc('ex5_data.txt') == (50, 643, 4298)
```

<!-- grade_id: cell-b9a248dad483a440, points: 0, grade: false, solution: true -->
```python
# Exercise 6 
def is_leap(year):
    # Complete this function to check if year is a leap year
    # Input: year
    # Output: True or False (Boolean)

    if (year % 4) == 0:
        if (year % 100) == 0:
            if (year % 400) == 0:
                return True
            else:
                return False
        else:
            return True
    else:
        return False
```

<!-- grade_id: cell-06fadd121addbb69, points: 4, grade: true, solution: false -->
```python
# Text cell for Exercise 6
assert is_leap(1800) == False
assert is_leap(1900) == False
assert is_leap(1600) == True
assert is_leap(2000) == True
```

<!-- grade_id: cell-39d284072956207c, points: 0, grade: false, solution: true -->
```python
# Exercise 7
def is_date_valid(month, day, year):
    # Complete this function to check if a data is valid, given month, day, and year.  
    # For example, 5/24/1962 is valid, but 9/31/2000 is not
    # Inputs: month, day, year
    # Output: True or False (Boolean)
    # Hint: Use is_leap() year function from previous exercise
    import datetime
    
    if(month==2):
        if(is_leap(year)):
            if(day<30):
                return True
            else:
                return False
        else:
            if(day<29):
                return True
            else:
                return False
    elif(month%2==1 and month!=7):
        if(day<31):
            return True
        else:
            return False
    elif(month%2==0 or month==7):
        if (day<31):
            return True
        else:
            return False

```

<!-- grade_id: cell-a8a1a1fa061bf04f, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 7
assert is_date_valid(5, 24, 1962) == True
assert is_date_valid(9, 31, 2000) == False
```

<!-- grade_id: cell-41144de972ffe16c, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 7 (hidden)
```

<!-- grade_id: cell-f81b2e6e7e373643, points: 0, grade: false, solution: true -->
```python
# Exercise 8
def day_number(month, day, year):
    # Complete this function to calculate the day_number given month, day, and year.
    # Information: The days of the year are often numbered from 1 through 365 (or 366).
    # This number can be computed in three steps using int arithmetic:
    # (a) - day_num = 31 * (month - 1) + day
    # (b) - if the month is after February subtract (4*(month)+23)//10
    # (c) - if it's a leap year and after February 29, add 1
    # Hint: First verify the input date is valid, return False if it is not valid; use is_date_valid
    # Hint: Also use the is_leap function
    # Inputs: month, day, year
    # Output: the day number or False (boolean) if the date is invalid. 
    
    if is_date_valid(month, day, year):
        
        if is_leap(year):
            
            if month > 2:
                print('after feb')
                return (31 * (month - 1) + day) - ((4*(month)+23)//10) + 1
            else:
                return (31 * (month - 1) + day)
        else:
            if month > 2:
                return (31 * (month - 1) + day) - ((4*(month)+23)//10)
            else:
                return (31 * (month - 1) + day)
    
    else:
        return False
    


```

<!-- grade_id: cell-17f6517a21d32010, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 8
assert day_number(9, 31, 2000) == False
assert day_number(2, 13, 2020) == 44
```

<!-- grade_id: cell-7bc040c625204f5d, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 8 (hidden)

```

<!-- grade_id: cell-c6da138b2c301741, points: 0, grade: false, solution: true -->
```python
# Exercise 9
def years_to_double_investment(apr):
    # Complete this function to use a while loop to determine how long it takes for an investment 
    # to double at a given interest rate. The input to this function, apr, is the annualized interest rate
    # and the output is the number of years it takes an investment to double. Note: The amount of the initial 
    # investment (principal) does not matter; you can use $1. 
    # Hint: principal is the amount of money being invested. 
    # apr is the annual percentage rate expressed as a decimal number.  
    # Relationship: value after one year is given by principal * (1+ apr)
    
    principal_amount = 1
    double_amount = principal_amount * 2
    counter = 0
    while principal_amount < double_amount:
        principal_amount = principal_amount * (1 + apr)
        counter = counter + 1
    
    return counter
```

<!-- grade_id: cell-72afae6aca21e6bc, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 9
assert years_to_double_investment(0.06) == 12
assert years_to_double_investment(0.1) == 8
```

<!-- grade_id: cell-56cac58d9d2e9e29, points: 0, grade: false, solution: true -->
```python
# Exercise 10
def stopping_time(n):
    # Complete this function to return the number of steps taken to reach 1 in
    # the Collatz sequence (https://en.wikipedia.org/wiki/Collatz_conjecture)

    counter = 0
    while n > 1:
        if n%2==0:
            n = n/2
        elif n%2==1:
            n = 3*n + 1
        counter = counter  + 1
    return counter

```

<!-- grade_id: cell-49b3568073d3a498, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 10
assert stopping_time(12) == 9
assert stopping_time(27) == 111
```

<!-- grade_id: cell-456f2e2dbffd8eb4, points: 2, grade: true, solution: false -->
```python
# Text cell for Exercise 10 (hidden)
```

<!-- grade_id: cell-f4384c44ad837566, points: 0, grade: false, solution: true -->
```python
# Exercise 11
def is_prime(n):
    # A positive whole number > 2 is prime if no number between 2 and sqrt(n)
    # (include) evenly divides n. Write a program that accepts a value of n as
    # input and determine if the value is prime. If n is not prime, your program should
    # return False (boolean) as soon as it finds a value that evenly divides n.
    # Hint: if number is 2, return False

    import math
    
    if n == 2:
        return False
    for i in range(2,n):
        if (n%i==0):
            return False
    return True
    

```

<!-- grade_id: cell-380e24833be40496, points: 4, grade: true, solution: false -->
```python
assert is_prime(2) == False
assert is_prime(3) == True
assert is_prime(5) == True
assert is_prime(25) == False
```

<!-- grade_id: cell-c76793df293d546e, points: 0, grade: false, solution: true -->
```python
# Exercise 12
def all_primes(n):
    # Complete this function to return all the primes as a list less than or equal to n
    # Input: n
    # Output: a list of numbers

    output = []
    for i in range(2,n+1):
        if is_prime(i):
            output.append(i)
    return (output)
```

```python
all_primes(22)
```

Output:
```
[3, 5, 7, 11, 13, 17, 19]
```

<!-- grade_id: cell-4f2b15776e494bec, points: 2, grade: true, solution: false -->
```python
assert all_primes(5) == [3, 5]
assert all_primes(25) == [3, 5, 7, 11, 13, 17, 19, 23]
```

<!-- grade_id: cell-422eed1aa7cba70e, points: 0, grade: false, solution: true -->
```python
# Exercise 13
def gcd(m,n):
    # Complete this function to determine the greatest common divisor (GCD). 
    # The GCD of two values can be computed using Euclid's algorithm. Starting with the values
    # m and n, we repeatedly apply the formula: n, m = m, n%m until m is 0. At this point, n is the GCD
    # of the original m and n. 
    # Inputs: m and n which are both natural numbers
    # Output: gcd

    if m == 0:
        #print(n)
        return n
    elif n == 0:
        #print(m)
        return m
    elif m > n:
        reminder = m % n
        quotient = round(m/n)
        m = n
        n = reminder
        return gcd(m,n)
    else:
        reminder = n % m
        quotient = round(n/m)
        n = m
        m = reminder
        return gcd(n,m)
    

```

<!-- grade_id: cell-6b64b0d91c4fc163, points: 2, grade: true, solution: false -->
```python
# Exercise 13 test cell
assert gcd(25,75) == 25
assert gcd(3,13) == 1
```

<!-- grade_id: cell-8d7e1f17c10c874a, points: 2, grade: true, solution: false -->
```python
# Exercise 13 test cell (hidden)
```

<!-- grade_id: cell-f0805833c03d7236, points: 0, grade: false, solution: true -->
```python
# Exercise 14 
def determine_min_max_average(filename):
    # Complete this function to read grades from a file and determine the student with the highest average 
    # test grades and the lowest average test grades. 
    # Input: filename
    # Output: a tuple containing four elements: name of student with highest average, their average, 
    # name of the student with the lowest test grade, and their average. Example ('Student1', 99.50, 'Student5', 65.50)
    # Hint: Round to two decimal places

    splited_line = []
    i=0
    highest_avg=99999.99
    highest_avg_student=0
    lowest_avg=00.00
    lowest_avg_student=0
    with open(filename, 'r') as f:
        for line in f:
            if i%2==0:
                splited_line.append(line.split(","))
            i = i + 1
        
        firstLine = 0
        for i in splited_line:
            #print((i))
            counter=0
            avg=0
            for j in i:
                if counter==0:
                    counter = 1
                    continue
                avg = avg + int(j)
            avg = round(avg/(len(i)-1),2)
            if firstLine == 0:
                highest_avg = avg
                lowest_avg = avg
                firstLine = 1
            elif firstLine > 0:
                if highest_avg < avg:
                    highest_avg = avg
                    highest_avg_student = i[0]
                if lowest_avg > avg:
                    lowest_avg = avg
                    lowest_avg_student = i[0]
                firstLine = 2
        return (highest_avg_student, highest_avg , lowest_avg_student, lowest_avg)
        
                
```

<!-- grade_id: cell-3c5bbcda431bee8f, points: 4, grade: true, solution: false -->
```python
# Exercise 14 test cell
assert determine_min_max_average('ex14_data.txt') == ('student733', 86.2, 'student202', 65.4)
```

<!-- grade_id: cell-ada31fff96cdef7c, points: 0, grade: false, solution: true -->
```python
# Ex15:Is a License Plate Valid?

def is_license_plate_valid(plate):

    # In a particular jurisdiction, older license plates consist of three uppercase
    # letters followed by three digits. When all of the license plates following
    # that pattern had been used, the format was changed to four digits followed by
    # three uppercase letters. 

    # Complete this function whose only input is a license plate and its output
    # is: 1) Older/Valid 2) Newer/Valid 3) Invalid
    # input: plate (str)
    # output: 'Older/Valid' or 'Newer/Valid' or 'Invalid'

    upper_char = 0
    number_char = 0
    
    if plate[0].isalpha() and plate[0].isupper():
        if (check_upper(plate) == True):
            return ('Older/Valid')
        else:
            return ('Invalid')
                
    elif plate[0].isnumeric():
        if check_num_new(plate) == True:
            return ('Newer/Valid')
        else:
            return ('Invalid')
    else:
        return ('Invalid')
        
def check_upper(str):
    counter = 0
    for i in str:
        if i.isalpha() :
            if i.isupper():
                counter = counter + 1
            else:
                return 'Invalid'
        else:
            if counter == 3 or counter == 4:
                return (check_num(counter, str))
                break
            else:
                return 'Invalid'
def check_num(index, str):
    counter = 0
    for i in range(index, len(str)):
        if str[i].isalpha():
            return 'Invalid'
        else:
            counter = counter + 1
    if counter == 3:
        return True
    else:
        return 'Invalid'
    
def check_num_new(str):
    counter = 0
    for i in str:
        if i.isnumeric():
            counter = counter + 1
        else:
            if counter < 4 or counter > 4:
                return 'Invalid'
    if counter == 4:
        return check_upper_new(counter, str)
    else:
        return 'Invalid'
    
def check_upper_new(index, str):
    counter = 0
    for i in range(index, len(str)):
        if str[i].isalpha() and str[i].isupper():
            counter = counter + 1
            
        else:
            return 'Invalid'
    if counter == 3:
        return True
    else:
        return 'Invalid'

```

```python
is_license_plate_valid('1243ACV')
```

Output:
```
'Newer/Valid'
```

<!-- grade_id: cell-a9212cfd8c67cd1b, points: 4, grade: true, solution: false -->
```python
# Exercise 15 test cell
assert is_license_plate_valid('ABC123') == 'Older/Valid'
assert is_license_plate_valid('GHE952') == 'Older/Valid'
assert is_license_plate_valid('1934ABT') == 'Newer/Valid'
assert is_license_plate_valid('bTR342') == 'Invalid'
```

<!-- grade_id: cell-ad96fa5eef4cb6eb, points: 4, grade: true, solution: false -->
```python
# Exercise 15 test cell (hidden)
```

<!-- grade_id: cell-b80e27c063f45d08, points: 0, grade: false, solution: true -->
```python
# Ex16: Magic dates

def is_magic_date(date):
    # A magic date is a date where the day multiplied by the month is equal 
    # to the two digit year. For example, June 10, 1960 is a magic date because
    # June is the sixth month, and 6 times 10 is 60, which is equal to the two
    # digit year. Complete this function to determine whether or not a date is
    # a magic date.

    # input: date (str -- month/day/year)
    # output: True or False (bool)

    month,day,year=date.split('/')
    if int(month)*int(day) == abs(int(year)) % 100:
        return True
    
```

```python
is_magic_date('6/10/1960')
```

Output:
```
True
```

<!-- grade_id: cell-01925812ea2e68b2, points: 2, grade: true, solution: false -->
```python
# Exercise 16 test cell
assert is_magic_date('6/10/1960') == True
assert is_magic_date('6/8/1948') == True
```

<!-- grade_id: cell-032789542acbe1e5, points: 0, grade: false, solution: true -->
```python
# Ex17
def remove_outliers(data, num_outliers):
    # When analyzing data collected as a part of a science experiment it 
    # may be desriable to remove the most extreme values before performing
    # other calculations. Complete this function which takes a list of
    # values and an non-negative integer, num_outliers, as its parameters.
    # The function should create a new copy of the list with the num_outliers 
    # largest elements and the num_outliers smallest elements removed. 
    # Then it should return teh new copy of the list as the function's only 
    # result. The order of the elements in the returned list does not have to
    # match the order of the lemetns in the original list.
    # input1: data (list)
    # input2: num_outliers (int)

    # output: list


    data.sort()
    data = data[:-2]
    i = num_outliers
    while i > 0:
        del data[0]
        i = i - 1
    return data

```

<!-- grade_id: cell-60da9757f2f02a8f, points: 4, grade: true, solution: false -->
```python
# Exercise 17 test cell
import random
random.seed(1234)
data = [random.randint(50, 150) for ele in range(100)]
data[45] = 1
data[46] = 2
data[90] = 250
data[34] = 300

result = [50, 50, 51, 52, 52, 52, 53, 53, 54, 
55, 55, 55, 55, 56, 58, 58, 59, 59, 59, 60, 61, 61, 
61, 62, 64, 64, 64, 68, 68, 68, 69, 69, 69, 70, 71, 
72, 73, 75, 77, 80, 81, 81, 84, 84, 84, 85, 88, 88, 89, 
92, 94, 94, 95, 95, 98, 103, 106, 108, 108, 109, 109, 109, 
110, 111, 111, 112, 113, 114, 115, 115, 117, 117, 119, 121, 
124, 124, 125, 126, 128, 129, 132, 132, 133, 134, 135, 135, 
135, 136, 138, 140, 141, 148, 148, 149, 149, 150]

assert remove_outliers(data, 2) == result

```

<!-- grade_id: cell-cd91dc245cfce54c, points: 4, grade: true, solution: false -->
```python
# Exercise 17 test cell (hidden)

```

<!-- grade_id: cell-ae1c664601ab69b7, points: 0, grade: false, solution: true -->
```python
# Ex18: Removing duplicates

def remove_duplicates(words):
    # Complete this function to remove duplicates from the words list using a loop
    # input: words (list)
    # output: a list without duplicates
    # MUST USE loop and NOT set!

    result = []
    for i in words:
        if i not in result:
            result.append(i)
    return result


```

<!-- grade_id: cell-e282aee0b81cce27, points: 2, grade: true, solution: false -->
```python
# Exercise 18 test cell
assert remove_duplicates([1,2,3,3,3,4,5,6,7,7,8]) == [1, 2, 3, 4, 5, 6, 7, 8]

```

<!-- grade_id: cell-0ee4d8aa6fbc4426, points: 2, grade: true, solution: false -->
```python
# Exercise 18 test cell (hidden)
```

<!-- grade_id: cell-658df9587965071c, points: 0, grade: false, solution: true -->
```python
# Ex19: List of proper divisors

def proper_divisors(n):
    # A proper divisor ofa  positive integer, n, is a positive integer less than n which divides
    # evenly into n. Complete this function to compute all the proper divisors of a positive
    # integer. The integer is passed to this function as the only parameter. The function will
    # return a list of containing all of the proper divisors as its only reult. 

    # input: n (int)
    # output: list

    result = []
    for i in range(1,n):
        if n % i == 0:
            result.append(i)
    return result
```

<!-- grade_id: cell-07080b1cba1b07a0, points: 2, grade: true, solution: false -->
```python
# Exercise 19 test cell
assert proper_divisors(28) == [1, 2, 4, 7, 14]
```

<!-- grade_id: cell-5c8065d0d450349a, points: 2, grade: true, solution: false -->
```python
# Exercise 19 test cell (hidden)
```

<!-- grade_id: cell-6967d2f5c29531b0, points: 0, grade: false, solution: true -->
```python
#Ex20: Perfect Numbers
def is_number_perfect(n):
    # An integer, n, is said to be perfect when the sum of all of the proper divisors 
    # of n is equal to n. For example, 28 is a perfect number because its proper divisors
    # are 1, 2, 4, 7, and 14 = 28 
    # Complete this function to determine if a the number a perfect number or not. 
    # input: n (int)
    # output: True or False (bool)


    list_of_divisors = proper_divisors(n)
    sum = 0
    for i in list_of_divisors:
        sum = sum + i
    if sum == n:
        return True
    else:
        return False
```

<!-- grade_id: cell-4611b4772a5abc36, points: 2, grade: true, solution: false -->
```python
# Exercise 20 test cell
assert is_number_perfect(28) == True
assert is_number_perfect(76) == False
```

<!-- grade_id: cell-0ba6c8ff7afe8991, points: 2, grade: true, solution: false -->
```python
# Exercise 20 test cell (hidden)

```

<!-- grade_id: cell-29c946e813372f5f, points: 0, grade: false, solution: true -->
```python
# Ex21: Line of Best Fit
def best_line(points):
    # Complete this function to determine the best line. 
    # https://www.varsitytutors.com/hotmath/hotmath_help/topics/line-of-best-fit
    # input: points (list of tuples contain x, y values)
    # output: (m, b) # round both values to two decimal places

    
    x_mean = 0
    y_mean = 0
    x_minus_mean = []
    y_minus_mean = []
    x_minus_mean_multi_y_minus_mean = []
    addition_x_minus_mean_multi_y_minus_mean = 0
    x_minus_mean_sq = []
    addition_x_minus_mean_sq = 0
    slope = 0.0
    y_intercept = 0
    for i in range(0, len(points)):
        x_mean += points[i][0]
        y_mean += points[i][1]
    x_mean = x_mean / len(points)
    y_mean = y_mean / len(points)
    for i in range(0, len(points)):
        x_minus_mean.append(points[i][0]-x_mean)
        y_minus_mean.append(points[i][1]-y_mean)
        x_minus_mean_multi_y_minus_mean.append(x_minus_mean[i] * y_minus_mean[i])
        addition_x_minus_mean_multi_y_minus_mean = addition_x_minus_mean_multi_y_minus_mean + x_minus_mean_multi_y_minus_mean[i]
        x_minus_mean_sq.append(x_minus_mean[i]**2)
        addition_x_minus_mean_sq = addition_x_minus_mean_sq + x_minus_mean_sq[i]
    slope = addition_x_minus_mean_multi_y_minus_mean/addition_x_minus_mean_sq
    y_intercept = y_mean - (slope*x_mean)
    #print(slope, y_intercept, y_mean, x_mean)
    return (round(slope,2), round(y_intercept,2))
```

```python
points = (
    (8, 3),
    (2, 10),
    (11, 3),
    (6, 6),
    (5, 8),
    (4, 12),
    (12, 1),
    (9, 4),
    (6, 9),
    (1, 14),
)

best_line(points)
```

Output:
```
(-1.11, 14.08)
```

<!-- grade_id: cell-13e17986e0d5080a, points: 4, grade: true, solution: false -->
```python
# Exercise 21 test cell
points = (
    (8, 3),
    (2, 10),
    (11, 3),
    (6, 6),
    (5, 8),
    (4, 12),
    (12, 1),
    (9, 4),
    (6, 9),
    (1, 14),
)

assert best_line(points) == (-1.11, 14.08)
```

<!-- grade_id: cell-06701dbdf8a04bd1, points: 4, grade: true, solution: false -->
```python
# Exercise 21 test cell (hidden)
```