Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tejas Yeshwantrao Wandile"
COLLABORATORS = ""
```

---

# Homework 8
Read the following link and complete this homework.
https://www.codemag.com/Article/1711091/Implementing-Machine-Learning-Using-Python-and-Scikit-learn

```python
import pandas as pd
```

# Step 1: Getting the Titanic Dataset

<!-- grade_id: cell-023793647b7f6e67, points: 0, grade: false, solution: true -->
```python
def step1():
    """
    Return a dataframe containing the Titantic dataset from the following URL
    # URL: https://gist.githubusercontent.com/mkzia/aa4f293661dba857b8c4459c0095ac95/raw/8075037f6f7689a1786405c1bc8ea9471d3aa9c3/train.csv

    """
    url = "https://gist.githubusercontent.com/mkzia/aa4f293661dba857b8c4459c0095ac95/raw/8075037f6f7689a1786405c1bc8ea9471d3aa9c3/train.csv"
    df=pd.read_csv(url)
    return df
    
```

<!-- grade_id: cell-ed6da60cf0a5fc56, points: 5, grade: true, solution: false -->
```python
df_from_func = step1()
df_from_file = pd.read_csv('step1_output.csv')
assert df_from_func.equals(df_from_file) == True

```

# Step 2: Clean data

<!-- grade_id: cell-09ac5c952374fa7f, points: 0, grade: false, solution: true -->
```python
def step2(df):
    """
    Modify df to drop the following columns:
    PassengerId
    Name
    Ticket
    Cabin
    Hint: Just pass all the columns to the .drop() method as an array
    """
    # YOUR CODE HERE
    df=df.drop(columns=['PassengerId', 'Name','Ticket','Cabin'])
    return df
    
```

<!-- grade_id: cell-2056e88ecd121be1, points: 5, grade: true, solution: false -->
```python
df_from_step1 = step1()
df_from_func = step2(df_from_step1)
df_from_file = pd.read_csv('step2_output.csv')
assert df_from_func.equals(df_from_file) == True
```

# Step 3: Drop NaNs and reindex
You want to reindex so your index does not have missing values after you drop the NaNs. Remember, index is used 
to access a row. Notice how many rows you dropped!

<!-- grade_id: cell-818d35e73cb1b82f, points: 0, grade: false, solution: true -->
```python
def step3(df):
    """
    Modify df to drop NaNs and reindex
    """
    df=df.dropna()
    df.reset_index(inplace=True)
    df=df.drop(columns=['index'])
    return df   
```

<!-- grade_id: cell-8c142adcb5eb7bed, points: 5, grade: true, solution: false -->
```python
df_from_step1 = step1()
df_from_step2 = step2(df_from_step1)
df_from_func = step3(df_from_step2)
df_from_file = pd.read_csv('step3_output.csv')
assert df_from_func.equals(df_from_file) == True
```

# Step 4: Encoding the Non-Numeric Fields

Encode text fields to numbers

<!-- grade_id: cell-4deb84aa50ed364e, points: 0, grade: false, solution: true -->
```python
def step4(df):
    """
    Modify df to encode Sex and Embarked to encoded values.
    """
    nums = {"Sex":     {"male": 1, "female": 0},
                   "Embarked":     {"S": 2, "C": 0, "Q": 1}}
    df = df.replace(nums)
    return df   
```

<!-- grade_id: cell-6c1849ea490b6f32, points: 5, grade: true, solution: false -->
```python
df_from_step1 = step1()
df_from_step2 = step2(df_from_step1)
df_from_step3 = step3(df_from_step2)
df_from_func = step4(df_from_step3)
df_from_file = pd.read_csv('step4_output.csv')
assert df_from_func.equals(df_from_file) == True
```

# Step 5: Making Fields Categorical

Turn values that are not continues values into categorical values

<!-- grade_id: cell-aa1b84a45200ba77, points: 0, grade: false, solution: true -->
```python
def step5(df):
    """
    Modify df to make Pclass, Sex, Embarked, and Survived a categorical field
    """
    # YOUR CODE HERE
    df['Pclass'] = pd.Categorical(df.Pclass)
    df['Sex'] = pd.Categorical(df.Sex)
    df['Embarked'] = pd.Categorical(df.Embarked)
    df['Survived'] = pd.Categorical(df.Survived)
    return df   
```

<!-- grade_id: cell-8226a758bb78aa66, points: 5, grade: true, solution: false -->
```python
df_from_step1 = step1()
df_from_step2 = step2(df_from_step1)
df_from_step3 = step3(df_from_step2)
df_from_step4 = step4(df_from_step3)
df_from_func = step5(df_from_step4)
df_from_file = pd.read_csv('step5_output.csv', skiprows=[0], header=None)

df_from_file.columns = ['', 1]
df_from_file.set_index('', inplace=True)
df_from_file.columns = ['']

df_from_func = pd.DataFrame(df_from_func.dtypes)
df_from_func.columns = ['']

assert df_from_func.equals(df_from_file) == True
```

# Step 6: 
1. Split dataframe into feature and label
2. Do train and test split; USE: random_state = 1
3. Use LogisticRegression() for classification
4. Return accuracy and confusion matrix

<!-- grade_id: cell-e94c61c148175fb4, points: 0, grade: false, solution: true -->
```python
def step6(df):
    """
    1. Split dataframe into feature and label
    2. Do train and test split; USE: random_state = 1
    4. Use LogisticRegression() for classification
    3. Return accuracy and confusion matrix
    
    Use  metrics.confusion_matrix to calculate the confusion matrix
    # https://towardsdatascience.com/understanding-confusion-matrix-a9ad42dcfd62
    # https://scikit-learn.org/stable/modules/generated/sklearn.metrics.confusion_matrix.html
    # IMPORTANT !!!! 
    # https://stackoverflow.com/questions/56078203/why-scikit-learn-confusion-matrix-is-reversed
    
    From the confusion matrix get TN, FP, FN, TP
    
    return --> accuracy, TN, FP, FN, TP; 
    Hint: round accuracy to 4 decimal places
    
    """
    from sklearn.model_selection import train_test_split
    from sklearn import linear_model
    from sklearn import metrics
    
    df_features = df.drop('Survived',axis = 1)
    df_label    = df['Survived']
    train_features,test_features,train_label,test_label = train_test_split(df_features,df_label,test_size = 0.25,random_state = 1,stratify = df["Survived"])
    
    logistic_regress_model = linear_model.LogisticRegression()
    logistic_regress_model.fit(X = train_features ,
    y = train_label)
    predicted = logistic_regress_model.predict(X=test_features)
    
    confusion_matrix = metrics.confusion_matrix(y_true = test_label,y_pred = predicted)  
    accuracy = round(logistic_regress_model.score(X = test_features, y = test_label),4)
    TN = confusion_matrix[0][0]
    FP = confusion_matrix[0][1]
    FN = confusion_matrix[1][0]
    TP = confusion_matrix[1][1]
    return accuracy, TN, FP, FN, TP 
```

<!-- grade_id: cell-527b3a7642f81e91, points: 5, grade: true, solution: false -->
```python
df_from_step1 = step1()
df_from_step2 = step2(df_from_step1)
df_from_step3 = step3(df_from_step2)
df_from_step4 = step4(df_from_step3)
df_from_step5 = step5(df_from_step4)
accuracy, TN, FP, FN, TP = step6(df_from_step5)

assert accuracy == 0.7978
assert TN == 90
assert FP == 16
assert FN == 20
assert TP == 52
```

```python

```